Free volume theory explains the unusual behavior of viscosity in a non-confluent tissue during morphogenesis

Author Response

The following is the authors’ response to the previous reviews.

We are pleased that Reviewers 1 and 3 have recommended that the revised paper be published.

Reviewer #2

For point A: Their preliminary simulation in 3D looks also nice, although it’s referenced in the discussion but not actually included in manuscript - I would advise adding it even under the mention of preliminary.

We appreciate the reviewer for liking our 3D results and suggesting to include them in the manuscript. However, these are preliminary results of our ongoing work. We are yet to establish the corresponding viscosity results quantitatively in the 3D simulations. Because the relationship between viscosity and relaxation time is not (always) linear in glass forming systems, we hesitate to report our results for publication. We hope to report the new results as part of a separate work.

For point B/C: I see some of the points of the authors - although not all of it made it in the main text. I still have some points that puzzle me. For instance, the authors mention that a single value of viscosity (from Green-Kubo) is ”valid for all time scales and amplitude”. This sounds very surprising to me for a complex fluid even at equilibrium: doesn’t it for instance assume linear response (hence small amplitudes)? Fast vs slow probing of a complex medium should also matter (see refs previously mentioned). Related to this, it’s not clear how can self-propulsion not matter if one would shear the system at a finite time scale, given past work on motility-driven unjamming and the mechanism of the authors from facilitation ( wouldn’t shearing at time scales larger vs smaller than the typical time for given cells to spontaneously rearrange from self-propulsion change drastically the effective complex modulus of the system?)

There might be a slight misunderstanding between the reviewer and us when

we say ‘single value of viscosity is valid for all time-scales and amplitude’. Let us explain this point more carefully. In our problem, we are studying the dynamics of a many body system which is undergoing Brownian dynamics where the fluctuation-dissipation theorem need not be valid (as the friction and the selfpropulsion noise strength are not related via Fluctuation-Dissipation Theorem). Now, for us to use the concepts of linear-response (which in the present study are the Green-Kubo relations for the transport coefficients in terms of timecorrelations functions), we need to show that the within the simulation time, the system has reached state that could be described using an “equilibrium” probability measure. This is the precise reason we calculated the ergodicity measure, which is a way to show that all the phase-space have been sampled uniformly under the given Brownian dynamics. This suggests (does not prove) that the system has attained a stationary probability measure (i.e, near equilibrium) for the value of self-propulsion used. Now for this value of self-propulsion, the Green-Kubo relations hold for ‘any time-scale of the simulations’ so that we can perform a time average over the trajectories of the particles (which is an alias of the stationary probability measure under the values of self-propulsion used). If we change the amplitude of the self-propulsion, we need to again compute the ergodicity measure and show the stationarity of the probability measure. If the system is ergodic with respect to the new self-propulsion, we can again use Green-Kubo for the simulations. Note that we will definitely get a different value of viscosity under the new self-propulsion as the shear-stresses generated will be different but the Green-Kubo holds. If the system is not ergodic, for the self-propulsion with the new amplitude, we cannot use Green-Kubo relations. Also a priori, one cannot say what is a large/small amplitude of self-propulsion because it has to be compared with the intrinsic energy scale, which is encoded in the energy function, which is difficult to say without explicit calculations.

This is what we meant when we said, ‘single value of viscosity is valid for all time-scales and amplitude’. It is valid for time-scales of the simulations for a given amplitude of self-propulsion only if the system is ergodic. Note that if the system is not ergodic, then the results of Ref. [14] (in the main text) could be questioned on theoretical grounds, because they were analyzed using 3 the equilibrium rigidity percolation theory. Nevertheless, the authors of Ref. [14] showed that equilibrium phase transition theory works in tissues. For these reasons, we have been, just like the Reviewer, puzzled that equilibrium ideas appear to be valid in the cell system. Additional theoretical work has to be done to clarify these links in tissues. Although this is not the last word, we hope this clarifies our view point.

For point D: I agree with the simplicity argument, although the added sentence from the discussion “Furthermore, the physics of the dynamics in glass forming materials does not change in systems with and without attractive forces” seems a bit strong given works like Lois et al., PRL, 2008 or Koeze et al, PRL, 2018 finding fundamentally different physics of jamming with or without adhesion. In the two cited papers the authors only consider equilibrium transitions in systems with attraction using computer simulations. Apparently, jamming properties depend on the strength of attraction. There are no attempts to characterize the dynamics, the focus of our work.

What we meant is that any universal relations, such as the Vogel-FulcherTammann relation, would still be valid. Of course, non-universal quantities such as glass transition temperature Tg or fragility will change. In our case, changing the adhesion strength would change ϕS, and the parameters in the VFT. However, our contention is that the overall finding that increase in viscosity followed by saturation is unlikely to change. We have added some clarifying statements in the manuscript to make this clear.

Changes from the earlier version to the revised version

Abstract

A recent experiment on zebrafish blastoderm morphogenesis showed that the viscosity (η) of a non-confluent embryonic tissue grows sharply until a critical cell packing fraction (ϕS). The increase in η up to ϕS is similar to the behavior observed in several glass forming materials, which suggests that the cell dynamics is sluggish or glass-like. Surprisingly, η is a constant above ϕS. To determine the mechanism of this unusual dependence of η on ϕ, we performed extensive simulations using an agent-based model of a dense non-confluent two-dimensional tissue. We show that polydispersity in the cell size, and the propensity of the cells to deform, results in the saturation of the available free area per cell beyond a critical packing fraction. Saturation in the free space not only explains the viscosity plateau above ϕS but also provides a relationship between equilibrium geometrical packing to the dramatic increase in the relaxation dynamics.

Introduction

There is great interest in characterizing the mechanical and dynamical properties of embryonic tissues because they regulate embryo development [1–5]. Measurements of bulk properties, such as viscosity and elastic modulus, and the dynamics of individual cells through imaging techniques, have been interpreted by adapting concepts developed to describe phase transitions, glass transition, and active matter [6–9]. Several experiments have shown that during embryo morphogenesis, material properties of the tissues change dramatically [10– 14]. Of relevance to our study is a remarkable finding that provided evidence that a phase transition (PT) occurs during zebrafish blastoderm morphogenesis, which was analyzed using rigidity percolation theory [14–17]. The authors also estimated the viscosity (η) of the blastoderm tissue using the micropipette aspiration technique [13, 18]. It was found that change in η is correlated with cell connectivity (⟨C⟩), rising sharply over a narrow range of ⟨C⟩. Surprisingly, a single geometrical quantity, the cell-cell contact topology controls both the rigidity PT as well as changes in η in this non-confluent tissue, thus linking equilibrium and transport properties.

Earlier version: Here, we focus on two pertinent questions that arise from the experiments on zebrafish blastoderm. First, experiments [14] showed that η increases as a function of the cell packing fraction (ϕ) till ϕ ≤ 0.87. The dependence of η on ϕ follows the well-known Vogel-FulcherTamman (VFT) law [19], which predicts that η grows monotonically with ϕ. The VFT law, which is commonly used to analyze the viscosity of a class of glass forming materials [20], is given by where ϕ is a constant. Surprisingly, for packing fractions, ϕ ≥ ϕS ≈ 0.90, η deviates from the VFT law and is independent of ϕ, which cannot be explained using conventional theories for glasses [8, 21]. Second, the experimental data [14] was interpreted using equilibrium rigidity percolation theory [15–17] for an embryonic tissue in which cells undergo random cell divisions. A priori it is unclear why equilibrium concepts should hold in zebrafish morphogenesis, which one would expect is controlled by non-equilibrium processes such as self-propulsion, growth and cell division.

Revised version: Here, we focus on two pertinent questions that arise from the experiments on zebrafish blastoderm. First, experiments [14] showed that η increases as a function of the cell packing fraction (ϕ) till ϕ ≤ 0.87. The dependence of η on ϕ follows the well-known Vogel-Fulcher-Tamman (VFT) law [19], which predicts that η grows monotonically with ϕ. The VFT law, which is commonly used to analyze the viscosity of a class of glass forming materials [20], is given by η expwhere ϕ0 is a constant. Surprisingly, for packing fractions, ϕ ≥ ϕS ≈ 0.90, η deviates from the VFT law and is independent of ϕ, which cannot be explained using conventional theories for glasses [8, 21]. Second, the experimental data [14] was interpreted using equilibrium rigidity percolation theory [15–17] for an embryonic tissue in which cells undergo random cell divisions. A priori it is unclear why equilibrium concepts should hold in zebrafish morphogenesis, which one would expect is controlled by non-equilibrium processes such as self-propulsion, growth and cell division.

We show that the two conundrums (saturation of η at high packing fractions and the use of equilibrium statistical mechanics in a growing system to explain phase transition) may be rationalized by (i) assuming that the interactions between the cells are soft, (ii) the cell sizes are highly heterogeneous (polydisperse), which is the case in zebrafish blastoderm. Using an agent-based (particle) simulation model of a two-dimensional non-confluent tissue, we explore the consequences of varying ϕ (see section Materials and Methods for the definition) of interacting self-propelled polydisperse soft cells, on η. The central results of our study are: (i) The calculated effective viscosity (for details see Appendix F), for the polydisperse cell system, shows that for ϕ ≤ ϕS ≈ 0.90 the increase in viscosity follows the VFT law. Just as in experiments, η is essentially independent of ϕ at high (≥ ϕS) packing fractions. The unusual dependence of η at ϕ ≥ ϕS is quantitatively explained using the notion of available free area fraction (ϕfree), which is the net void space that can be explored by the cells when they are jammed. At high densities, a given cell requires free space in order to move. The free area is created by movement of the neighboring cells into the available void space. One would intuitively expect that the ϕfree should decrease with increasing packing fractions, due to cell jamming, which should slow down the overall dynamics. Indeed, we find that ϕfree decreases with increasing packing fraction (ϕ) until ϕS. The simulations show that when ϕ exceeds ϕS, the free area ϕfree saturates because the soft cells (characterized by “soft deformable disks”) can overlap with each other, resulting in the collective dynamics of cells becoming independent of ϕ for ϕ ≥ ϕS. As a consequence η saturates at high ϕ. (iii) Cells whose sizes are comparable to the available free area move almost like a particle in a liquid. The motility of small sized cells facilitates adjacent cells to move through multi-cell rearrangement even in a highly jammed environment. The facilitation mechanism, invoked in glassy systems [22] allows large cells to move with low mobility. A cascade of such facilitation processes enable all the cells to remain dynamic even above the onset packing fraction of the phase transition. (iv) We find that the relaxation time does not depend on the waiting time for measurements even in the regime where viscosity saturates. In other words, there is no evidence of aging even in the regime where viscosity saturates. Strikingly, the tissue exhibits ergodic [23] behavior at all densities. The cell-based simulations, which reproduce the salient experimental features, may be described using equilibrium statistical mechanics, thus providing credence to the use of cell contact mechanics to describe both rigidity PT and dynamics in an evolving non-confluent tissue [14].

Results

Experimental results

We first describe the experimental observations, which serve as the basis for carrying out the agent-based simulations. Fig. 1(A) shows the bright-field images of distinct stages during zebrafish morphogenesis. A 2D section of zebrafish blastoderm (Fig. 1 (B)) shows that there is considerable dispersion in cell sizes. The statistical properties of the cell sizes are shown in Fig. 8 (D). Fig. 1 (C) shows that η increases sharply over a narrow ϕ range, and saturates when ϕ exceeds ϕS ≈ 0.90.

Structure and viscosity of non-confluent tissues:

Earlier version: (A) Bright-field single-plane images of an exemplary embryo of zebrafish before (t = −60 min), at the onset (t = 0 min), and after blastoderm spreading (t = 60 min). (B) Snapshot of 2D confocal sections at the 1st–2nd deep-cell layer of the blastoderm at t = 60 min. (A) and (B) are taken from [14]. (C) Viscosity η of zebrafish blastoderm as a function of ϕ in a log-linear scale using the data from [14]. The dashed line is the fit to VFT equation. Note that η does not change significantly beyond ϕ ≥ 0.87. (D) A typical snapshot taken from cell-based simulations for ϕ = 0.93. Cells are colored according to their radii (in μm) (color bar shown on the right). (E) The pair correlation function, g(r), as a function of r for ϕ = 0.93. The vertical dashed line is the position of the first peak (rmax = 17.0 μm). The pair correlation function does not exhibit signs of long-range order. Scale bars in (A) is 100 μm and in (B) is 50 μm.

Revised version: (A) Bright-field single-plane images of an exemplary embryo of zebrafish before (t = −60 min), at the onset (t = 0 min), and after blastoderm spreading (t = 60 min). (B) Snapshot of 2D confocal sections at the 1st–2nd deep-cell layer of the blastoderm at t = 60 min. (A) and (B) are taken from [14]. (C) Viscosity η of zebrafish blastoderm as a function of ϕ in a log-linear scale using the data from [14]. The dashed line is the fit to VFT equation. Note that η does not change significantly beyond ϕ ≥ 0.87. (D) A typical snapshot taken from cell-based simulations for ϕ = 0.93. Cells are colored according to their radii (in μm) (color bar shown on the right). (E) The pair correlation function, g(r), as a function of r for ϕ = 0.93. The vertical dashed line is the position of the first peak (rmax = 17.0 μm). The pair correlation function does not exhibit signs of long-range order. Scale bars in (A) is 100 μm and in (B) is 50 μm. range order. Thus, the polydisperse cell system exhibits liquid-like structure even at the high ϕ.

To account for the results in Fig. 1 (C), we first simulated a mono-disperse system in which all the cells have identical radius (R = 8.5 μm). Because the system crystallizes (Fig. 8 (A) & (B)), we concluded that the dynamics observed in experiments cannot be explained using this model. A 1:1 binary mixture of cells, with different radii gives glass-like behavior for all ϕ, with the relaxation time τα as well as the effective viscosity (defined in Eq. 1) following the VFT behavior (see Appendix B).

Polydispersity and cell-cell interactions

In typical cell tissues, and zebrafish in particular, there is a great dispersion in the cell sizes, which vary in a single tissue by a factor of ∼ 5 − 6 [14] (Fig. 1 (B) and Fig. 8 (D)). In addition, the elastic forces characterizing cell-cell interactions are soft, which implies that the cells can overlap, with rij −(Ri +Rj) < 0 when they are jammed ((Fig. 1 (B) and (Fig. 1 (D)). Thus, both polydispersity as well as soft interactions between the cells must control the relaxation dynamics. To test this proposition, we simulated a highly polydisperse system (PDs) in which the cell sizes vary by a factor of ∼ 8 (Fig. 1 (D) and Fig. 8 (E)).

Earlier version: A simulation snapshot (Fig. 1 (D)) for ϕ = 0.93 shows that different sized cells are well mixed. In other words, the cells do not phase separate. The structure of the tissue can be described using the pair correlation function, , where is the number density, δ is the Dirac delta function, is the position of the ith cell, and the angular bracket ⟨⟩ denotes an average over different ensembles. The g(r) function (Fig. 1 (E)) has a peak around r ∼ 17µm, which is approximately the average diameter of the cells. The absence of peaks in g(r) beyond the second one, suggests there is no long range order. Thus, the polydisperse cell system exhibits liquid-like structure even at the high ϕ.

Revised version: A simulation snapshot (Fig. 1 (D)) for ϕ = 0.93 shows that different sized cells are well mixed. In other words, the cells do not phase separate. The structure of the tissue can be described using the pair correlation function, , where is the number density, δ is the Dirac delta function, is the position of the ith cell, and the angular bracket ⟨⟩ denotes an average over different ensembles. The g(r) function (Fig. 1 (E)) has a peak around r ∼ 17μm, which is approximately the average diameter of the cells. The absence of peaks in g(r) beyond the second one, suggests there is no long

Effective shear viscosity() as a function of ϕ

Earlier version: A fit of the experimental data for η using the VFT [24, 25] relation in the range ϕ ≤ 0.87 (Fig. 1 (C)) yields ϕ0 ≈ 0.95 and D ≈ 0.51 [19]. The VFT equation for cells, which is related to the Doolittle equation [26] for fluidity that is based on free space available for motion in an amorphous system [27, 28], is , where D is the apparent activation energy. In order to compare with experiments, we calculated an effective shear viscosity for the polydisperse system using a Green-Kubo-type relation [29], The stress tensor Pµν(t) in the above equation is, where µ, ν ∈ (x, y) are the Cartesian components of coordinates, is the force between ith and jth cells and A is the area of the simulation box. Note that should be viewed as a proxy for shear viscosity because it does not contain the kinetic term and the factor is not included in Eqn. (1) because temperature is not a relevant variable in the highly over-damped model for cells considered here.

Revised version: A fit of the experimental data for η using the VFT [24, 25] relation in the range ϕ ≤ 0.87 (Fig. 1 (C)) yields ϕ0 ≈ 0.95 and D ≈ 0.51 [19]. The VFT equation for cells, which is related to the Doolittle equation [26] for fluidity that is based on free space available for motion in an amorphous system [27, 28], is η = η exp, where D is the apparent activation energy. In order to compare with experiments, we calculated an effective shear viscosity() for the polydisperse system using a Green-Kubo-type relation [29], The stress tensor Pμν(t) in the above equation is, where μ, ν ∈ (x, y) are the Cartesian components of coordinates, is the force between ith and jth cells and A is the area of the simulation box. Note that should be viewed as a proxy for shear viscosity because it does not contain the kinetic term and the factor is not included in Eqn. (1) because temperature is not a relevant variable in the highly over-damped model for cells considered here.

Earlier version: Plot of as a function of ϕ in Fig. 2 (A), shows qualitatively the same behavior as the estimate of viscosity (using dimensional arguments) made in experiments. Two features about Fig. 1 (C) and Fig. 2 (A) are worth noting. (i) Both simulations and experiments show that up to follows the VFT relation with ϕ0 ∼ 0.94 and D ∼ 0.5. More importantly is independent of ϕ when ϕ > 0.90. (ii) The values of ϕ0 and D obtained by fitting the experimental estimate of η to the VFT equation and simulation results are almost identical. Moreover, the onset of the plateau packing fraction in simulations and experiments occur at the same value (ϕS ∼ 0.90). The overall agreement with experiments is remarkable given that the model was not created to mimic the zebrafish tissue.

Revised version: Plot of as a function of ϕ in Fig. 2 (A), shows qualitatively the same behavior as the estimate of viscosity (using dimensional arguments) made in experiments. Two features about Fig. 1 (C) and Fig. 2 (A) are worth noting. (i) Both simulations and experiments show that up to ϕ ≈ 0.90, follows the VFT relation with ϕ0 ∼ 0.94 and D ∼ 0.5. More importantly is independent of ϕ when ϕ > 0.90. (ii) The values of ϕ0 and D obtained by fitting the experimental estimate of η to the VFT equation and simulation results are almost identical. Moreover, the onset of the plateau packing fraction in simulations and experiments occur at the same value (ϕS ∼ 0.90). The overall agreement with experiments is remarkable given that the model was not created to mimic the zebrafish tissue.

Saturation in viscosity and relaxation time:

(A) Effective viscosity as a function of ϕ, with the solid line being the fit to VFT equation. The inset shows at high ϕ. The dashed line in the inset is the expected behavior assuming that the VFT relation holds at all ϕ. (B) The self-intermediate scattering function Fs(q, t) as a function of t for 0.70 ≤ ϕ ≤ 0.905. The dashed line corresponds to . (C) A similar plot for ϕ > 0.905. (D) The logarithm of the relaxation time τα(s) as a function of ϕ. The VFT fit is given by the dashed line. The inset shows a zoomed-in view for ϕ ≥ ϕS.

Earlier version: To provide additional insights into the dynamics, we calculated the isotropic selfintermediate scattering function, Fs(q, t), where is the wave vector, and is the position of a cell at time t. The degree of dynamic correlation between two cells can be inferred from the decay of Fs(q, t). The angle bracket ⟨…⟩ is an average over different time origins and different trajectories. We chose , where rmax is the position of the first peak in g(r) between all cells (see Fig. 1 (E)). The relaxation time τα is calculated using .

Revised version: To provide additional insights into the dynamics, we calculated the isotropic self-intermediate scattering function, Fs(q, t), Where is the wave vector, and is the position of a cell at time t. The degree of dynamic correlation between two cells can be inferred from the decay of Fs(q, t). The angle bracket ⟨…⟩ is an average over different time origins and different trajectories. We chose where rmax is the position of the first peak in g(r) between all cells (see Fig. 1 (E)). The relaxation time τα is calculated using .

From Figs. 2 (B) & (C), which show Fs(q, t) as a function of t for various ϕ, it is clear that the dynamics become sluggish as ϕ increases. The relaxation profiles exhibit a two step decay with a plateau in the intermediate time scales. The dynamics continues to slow down dramatically until ϕ ≤ 0.90. Surprisingly, the increase in the duration of the plateau in Fs(q, t) ceases when ϕ exceeds ≈ 0.90 (Fig. 2 (C)), a puzzling finding that is also reflected in the dependence of τα on ϕ in Fig. 2 (D). The relaxation time increases dramatically, following the VFT relation, till ϕ ≈ 0.90, and subsequently reaches a plateau (see the inset in Fig. 2 (D)). If the VFT relation continued to hold for all ϕ, as in glasses or in binary mixture of 2D cells (see Appendix B), then the fit yields ϕ0 ≈ 0.95 and D ≈ 0.50. However, the simulations show that τα is nearly a constant when ϕ exceeds 0.90. We should note that the behavior in Fig. 2 (D) differs from the dependence of τα on ϕ for 2D monodisperse polymer rings, used as a model for soft colloids. Simulations [30] showed τα increases till a critical ϕS but it decreases substantially beyond ϕS with no saturation.

Relaxation dynamics of individual cells

Plot of τα as a function of the radius of cells Ri (Fig. 3 (A)) shows nearly eight orders of magnitude change. The size dependence of τα on ϕ is striking. That τα should increase for large sized cells (see the data beyond the vertical dashed line in Fig. 3 (A)) is not unexpected. However, even when cell sizes increase beyond Ri = 4.25 μm, the dispersion in τα is substantial, especially when ϕ exceeds ϕS. The relaxation times for cells with Ri < 4.25 μm are relatively short even though the system as a whole is jammed. For ϕ ≥ 0.90, τα for small sized cells have a weak dependence on ϕ. Although τα for cells with radius < 4 μm is short, it is clear that for a given ϕ (for example ϕ = 0.93) the variations in τα is substantial. In contrast, τα’s for larger cells (R ≥ 7μm) are substantially large, possibly exceeding the typical cell division time in experiments. In what follows, we interpret these results in terms of available free area ⟨Afree⟩ for cells. The smaller sized cells have the largest (RS is the radius of the small cell).

Spectrum of relaxation times:

(A) Scatter plot of relaxation times τα(s) as a function of cell radius. From top to bottom, the plot corresponds to decreasing ϕ. The vertical dashed line is for Ri = 4.25 μm, beyond which the τα changes sharply at high packing fractions. (B) Histogram P (ln(τα)) as a function of ln(τα). Beyond ϕ = 0.90 (ϕS), the histogram peaks do not shift substantially towards a high τα values. (C) For ϕ ≤ ϕS P (ln(τα)) (scaled by P max(ln(τα))) falls on a master curve, as described in the main text. (D) Same as (C) except the results are for ϕ > 0.90. The data deviates from the Gaussian fit, shown by the dashed line.

The effect of jamming on the dramatic increase in τα occurs near Ri ≈ 4.5 μm, which is comparable to the length scale of short range interactions. For ϕ ≤ 0.90, τα increases as the cell size increases. However, at higher packing fractions, even cells of similar sizes show substantial variations in τα, which change by almost 3 − 4 orders of magnitude (see the data around the vertical dashed line for ϕ ≥ 0.915 in Fig. 3 (A)). This is a consequence of large variations in the local density (Fig. 14). Some of the similar-sized cells are trapped in the jammed environment, whereas others are in less crowded regions (see Fig. 14). The spread in τα increases dramatically for ϕ > ϕS (≈ 0.90) and effectively overlap with each other. This is vividly illustrated in the the histogram, P(log(τα)) shown in Fig. 3 (B). For ϕ < ϕs, the peak in P(log(τα)) monotonically shifts to higher log(τα) values. In contrast, when ϕ exceeds ϕS there is overlap in P(log(τα)), which is reflected in the saturation of and τα. There are cells (typically with small sizes) that move faster even in a highly jammed environment (see Fig. 12 (C) and Fig. 13). The motions of the fast-moving cells change the local environment, which effectively facilitates the bigger cells to move in a crowded environment (see Fig. 12 (D) and Fig. 13 and and Movie 1 (ϕ = 0.92 > ϕS) and Movie 2 (ϕ = 0.90 = ϕS)). In contrast, for ϕ = 0.85 < ϕS small and large sized cells move without hindrance because of adequate availability of free area (Movie 3). The videos vividly illustrate the large scale facilitated rearrangements that enable the large sized cells to move.

The dependence of τα on ϕ for ϕ ≤ ϕS (Fig. 2 (D)) implies that the polydisperse cell systems behave as a soft glass in this regime. On theoretical grounds, it was predicted that in glass-forming systems [8]. Remarkably, we found that this prediction is valid in the polydisperse cell system (Fig. 3 (C)). However, above ϕS the predicted relation is not satisfied (see Fig. 3 (D)).

Available free area explains viscosity saturation at high ϕ

We explain the saturation in the viscosity by calculating the available free area per cell, as ϕ increases. In a hard disk system, one would expect that the free area would decrease monotonically with ϕ until it is fully jammed at the close packing fraction (∼ 0.84 [31, 32]). Because the cells are modeled as soft deformable disks, they could overlap with each other even when fully jammed. Therefore, the region where cells overlap create free area in the immediate neighborhood. The extent of overlap (hij) is reflected in distribution P(hij). The width in P(hij) increases with ϕ, and the peak shifts to higher values of hij (Fig. 4 (A)). The mean, ⟨hij⟩ increases with ϕ (Fig. 4 (B)). Thus, even if the cells are highly jammed at ϕ ≈ ϕS, free area is available because of an increase in the overlap between cells (see Fig. 5). When ϕ exceeds ϕS, the mobility of small sized cells facilitates the larger cells to move, as is assumed in the free volume theory of polymer glasses [28, 33–35]. As a result of the motion of small cells, a void is temporarily created, which allows other (possibly large) cells to move. In addition to the release of space, the cells can also interpenetrate (Figs. 4 (A) & (B)). If hij increases, as is the case when the extent of compression increases (Figs. 4 (A)), the available space for nearby cells would also increase. This effect is expected to occur with high probability at ϕS and beyond, resulting in high overlap between the cells. These arguments suggest that the combined effect of polydispersity and cell-cell overlap creates, via the self-propulsion of cells, additional free area that drives larger cells to move even under jammed conditions.

Density dependent cell-cell overlap :

(A) Probability of overlap (hij) between two cells, P (hij), for various ϕ values. The peak in the distribution function shifts to higher values as ϕ increases. (B) Mean ⟨hij⟩ = dhijP (hij) as a function of ϕ. Inset shows a pictorial illustration of h12 between two cells with radii R1 and R2 at a distance r12.

Changes in free area fraction with ϕ:

(A) Voronoi tessellation of cells for ϕ = 0.93 for a single realization. The orange circles represent actual cell sizes. The blue polygons show the Voronoi cell size. (B) Distribution of Voronoi cell size A as a function of ϕ. (C) Mean Voronoi cell size ⟨A⟩ as a function of ϕ. A zoomed-in view for ϕ > 0.860 is shown in the inset. (D) Distribution of free area P (Afree) for all ϕ. The vertical blue dashed line shows that the maximum in the distribution is at Afree ∼ 50μm2. (E) Free area fraction ϕfree as a function of ϕ. Note that ϕfree saturates beyond ϕ = 0.90. An expanded view of the saturated region is shown in the right panel of (E).

In order to quantify the physical picture given above, we calculated an effective area for each cell by first calculating Voronoi cell area A. A plot for Voronoi tessellation is presented in Fig. 5 (A) for ϕ = 0.93, and the histogram of A is shown in Fig. 5 (B). As ϕ increases, the distribution shifts towards lower Voronoi cell size ⟨A⟩. The mean Voronoi cell size ⟨A⟩ as a function of ϕ in Fig. 5 (C) shows ⟨A⟩ decreases as ϕ is increased. As cells interpenetrate, the Voronoi cell size will be smaller than the actual cell size in many instances (Fig. 5 (A)). To demonstrate this quantitatively, we calculated . The value of Afree could be negative if the overlap between neighbouring cells is substantial; Afree is positive only when the Voronoi cell size is greater than the actual cell size. Positive Afree is an estimate of the available free area. The histograms of Afree for all the packing fractions in Fig. 5 (D) show that the distributions saturate beyond ϕ = 0.90. All the distributions have a substantial region in which Afree is negative. The negative value of Afree increases with increasing ϕ, which implies that the amount of interpenetration between cells increases.

Because of the overlap between the cells, the available free area fraction ϕfree, is higher than the expected free area fraction (1.0 − ϕ) for all ϕ. We define an effective free area fraction ϕfree as, where Np is the number of positive free area in jth snapshots, Nt is the total number of snapshots, Abox is the simulation box area and is the positive free area of ith cell in jth snapshot. The calculated ϕfree, plotted as a function of ϕ in Fig. 5 (E), shows that ϕfree decreases with ϕ until ϕ = 0.90, and then it saturates near a value ϕfree ≈ 0.22 (see the right panel in Fig. 5 (E)). Thus, the saturation in as a function of ϕ is explained by the free area picture, which arises due to combined effect of the size variations and the ability of cells to overlap.

Earlier version: F. Aging does not explain viscosity saturation

Revised version: Aging does not explain viscosity saturation

Earlier version: Our main result, which we explain by adopting the free volume theory developed in the context of glasses [28, 33–35], is that above a critical packing fraction ϕS ∼ 0.90 the viscosity saturates. Relaxation time, τα, measured using dynamic light scattering, in nearly monodisperse microgel poly(Nisopropylacrylamide) (PNiPAM) [36] was found to depend only weakly on the volume fraction (3D), if ϕV, exceeds a critical value. It was suggested that the near saturation of τα at high ϕV is due to aging, which is a non-equilibrium effect. If saturation in viscosity and relaxation time in the embryonic tissue at high ϕ is due to aging then τα should increase sharply as the waiting time, τω, is lengthened. We wondered if aging could explain the observed saturation of η in the embryonic tissue above ϕS. If aging causes the plateau in the tissue dynamics, then η or τα should be an increasing function of the waiting time, τω. To test the effect of τω on τα, we calculated the self-intermediate scattering function Fs(q, t+τω) as a function of t by varying τω over three orders of magnitude at ϕ = 0.92 (Fig. 6 (A)). There is literally no change in Fs(q, t + τω) over the entire range of tω. We conclude that, τα, extracted from Fs(q, t+τω) is independent of τω. The variations in τα (Fig. 6 (B)), with respect to τω, is significantly smaller than the errors in the simulation. Thus, the saturation in η or τα when ϕ > ϕS is not a consequence of aging.

Revised version: Our main result, which we explain by adopting the free volume theory developed in the context of glasses [28, 33–35], is that above a critical packing fraction ϕS ∼ 0.90 the viscosity saturates. Relaxation time, τα, measured using dynamic light scattering, in nearly monodisperse microgel poly(N-isopropylacrylamide) (PNiPAM) [36] was found to depend only weakly on the volume fraction (3D), if ϕV, exceeds a critical value. It was suggested that the near saturation of τα at high ϕV is due to aging, which is a non-equilibrium effect. If saturation in viscosity and relaxation time in the embryonic tissue at high ϕ is due to aging then τα should increase sharply as the waiting time, τω, is lengthened. We wondered if aging could explain the observed saturation of η in the embryonic tissue above ϕS. If aging causes the plateau in the tissue dynamics, then η or τα should be an increasing function of the waiting time, τω. To test the effect of τω on τα, we calculated the self-intermediate scattering function Fs(q, t+τω) as a function of t by varying τω over three orders of magnitude at ϕ = 0.92 (Fig. 6 (A)). There is literally no change in Fs(q, t + τω) over the entire range of tω. We conclude that, τα, extracted from Fs(q, t+τω) is independent of τω. The variations in τα (Fig. 6 (B)), with respect to τω, is significantly smaller than the errors in the simulation. Thus, the saturation in η or τα when ϕ > ϕS is not a consequence of aging.

Relaxation in the polydisperse cell system is independent of the waiting time:

(A) Fs(q, t) for ϕ = 0.92 at different waiting times (τω = 106(s)). Regardless of the value of τω, all the Fs(q, t) curves collapse onto a master curve. (B) Relaxation time, ln(τα), as a function of τω. Over a 3 orders of magnitude change in tω, the variation in relaxation times is less than the sample-to-sample fluctuations, as shown by the error bar.

Earlier version: There are two implications related to the absence of aging in the dynamics of the nonconfluent embryonic tissues. (i) Although active forces drive the dynamics of the cells, as they presumably do in reality, the cell collectives can be treated as being near equilibrium, justifying the use of Green-Kubo relation to calculate η. (ii) Parenthetically, we note that the absence of significant non-equilibrium effects, even though Zebrafish is a living system, further justifies the use of equilibrium rigidity percolation theory to analyze the experimental data [14].

Revised version: There are two implications related to the absence of aging in the dynamics of the non-confluent embryonic tissues. (i) Although active forces drive the dynamics of the cells, as they presumably do in reality, the cell collectives can be treated as being near equilibrium, justifying the use of Green-Kubo relation to calculate η. (ii) Parenthetically, we note that the absence of significant non-equilibrium effects, even though Zebrafish is a living system, further justifies the use of equilibrium rigidity percolation theory to analyze the experimental data [14].

Discussion

Extensive computer simulations of a two-dimensional dense tissue using a particle-based model of soft deformable cells with active self-propulsion have successfully reproduced the dynamical behavior observed in the blastoderm tissue of zebrafish. The dependence of viscosity (η) and relaxation time (τα) (before the saturation) is well fit by the VFT equation. The value of ϕ0 obtained from simulations, ϕ0 ∼ 0.95, is close to ϕ0 ∼ 0.94 extracted by fitting the experimental data to the VFT equation. Thus, the dynamics for ϕ ≤ ϕS resembles the behavior expected for glass forming systems. Remarkably, the dependence of η on ϕ over the entire range (VFT regime followed by a plateau) may be understood using available free area picture with essentially a single parameter, an idea that was proposed nearly 70 years ago. We discovered that polydispersity as well as the ease of deformation of the cells that creates free area under high jamming conditions, is the mechanism that explains viscosity saturation at high cell densities. The mechanism suggested here is an important step that links equilibrium phase transition to dynamics during zebrafish development [37].

One could legitimately wonder if the extent of polydispersity (PD) used in the soft discs simulations, which seems substantial, is needed to recapitulate the observed dependence of η on ϕ. Furthermore, such large values of PD may not represent biological tissues. Although the choice of PD was made in part by the two-dimensional projection of area reported in experiments [14], it is expected that PD values have to be less in three dimensions. We performed preliminary simulations in three dimensions with considerably reduced PD, and calculated the dependence of relaxation time (τα) as a function of ϕ. The results show that τα does indeed saturate at high volume fractions.

Earlier version: The proposed model neglects adhesive interactions between cells, which of course is not unimportant. It is crucial to wonder if the proposed mechanism would change if cell-cell adhesion is taken into account. We wanted to create the simplest model to explain the experimental data. We do think that realistic values of adhesion strength does not significantly alter the forces between cells [38]. Thus, we expect a similar mechanism. Furthermore, the physics of the dynamics in glass forming materials does not change in systems with and without attractive forces [8]. For these reasons, in the first pass we neglected adhesion, whose effects have to be investigated in the future.

Revised version: The proposed model neglects adhesive interactions between cells, which of course is not unimportant. It is crucial to wonder if the proposed mechanism would change if cell-cell adhesion is taken into account. We wanted to create the simplest model to explain the ex-perimental data. We do think that realistic values of adhesion strength does not significantly alter the forces between cells [38]. Thus, we expect a similar mechanism. Furthermore, the physics of the dynamics in glass forming materials does not change in systems with and without attractive forces [8]. Universal behavior, such as Vogel-Fucher-Tammann relation, is valid for a broad class of unrelated materials (see Fig.1 in [20]). Needless to say, non-universal quantities such as glass transition temperature Tg or effective free energy barriers for relaxation will change. In our case, we expect that changing the adhesion strength, within a reasonable range, would change ϕS without qualitatively altering the dependence of η on ϕ. For these reasons, in the first pass we neglected adhesion, whose effects have to be investigated in the future.

In the physical considerations leading to Eq. 6, the random activity term (μ) plays an important role. Is it possible to create a passive model by maintaining the system at a finite temperature using stochastic noise with μ = 0, which would show the observed viscosity behavior? First, in such a system of stochastic equations, the coefficient of noise (a diffusion constant) would be related to γi in Eq.6 through fluctuation dissipation theorem (FDT). Thus, only γi can be varied. In contrast, in Eq.6 the two parameters (γi and μ) maybe independently changed, which implies that the two sets of stochastic equations of motion are not equivalent. Second, the passive system describes particles that interact by soft Hertz potential. In analogy with systems in which the particles interact with harmonic potential [39], we expect that the passive model would form a glass in which the viscosity would follow the VFT law.

We find it surprising that the calculation of viscosity using linear response theory (valid for systems close to equilibrium), and the link to free area quantitatively explain the simulation results and by implication the experimental data for a living and growing tissue. The calculation of free area of the cells is based on the geometrical effects of packing, which in turn is determined by cell to cell contact topology. These considerations, that are firmly established here, explain why equilibrium phase transitions are related to steep increase in viscosity [8] as the packing fraction changes over a narrow range. The absence of aging suggests that, although a large number of cell divisions occur, they must be essentially independent, thus allowing the cells to reach local equilibrium.